```python
# Initialize Otter
import otter
grader = otter.Notebook("hw8.ipynb")
```

# CPSC 330 - Applied Machine Learning 

## Homework 8: Time series
**Due date: See the [Calendar](https://calendar.google.com/calendar/u/0/embed?src=7a04205ae91b85e82ebc74daddbf2933c6b6723b81abb966f0e69c66a996c43b@group.calendar.google.com&ctz=America/Vancouver&pli=1).**

## Imports

```python
from hashlib import sha1

import numpy as np
import pandas as pd
import matplotlib.pyplot as plt

from sklearn.impute import SimpleImputer
from sklearn.compose import ColumnTransformer
from sklearn.pipeline import Pipeline
from sklearn.preprocessing import StandardScaler, OrdinalEncoder, OneHotEncoder

from sklearn.linear_model import Ridge
from sklearn.ensemble import RandomForestRegressor

from sklearn.metrics import r2_score
```

## Submission instructions
<hr>
rubric={points:4}

You will receive marks for correctly submitting this assignment. 

Follow the homework instructions below and at the end of this file. 
There are also detailed [homework submission instructions](https://github.com/UBC-CS/cpsc330-2023s/blob/main/docs/homework_instructions.md) on github.

- **You may work on this assignment in a group (group size <= 4) and submit your assignment as a group.** 
- Below are some instructions on working as a group.  
    - The maximum group size is 4. 
    - You can choose your own group members. 
    - Use group work as an opportunity to collaborate and learn new things from each other. 
    - Be respectful to each other and make sure you understand all the concepts in the assignment well. 
    - It's your responsibility to make sure that the assignment is submitted by one of the group members before the deadline. [Here](https://help.gradescope.com/article/m5qz2xsnjy-student-add-group-members) are some instructions on adding group members in Gradescope.  
- Upload the .ipynb file to Gradescope.
- **If the .ipynb file is too big or doesn't render on Gradescope for some reason, also upload a pdf or html in addition to the .ipynb.** 
- Make sure that your plots/output are rendered properly in Gradescope.

<br><br>

## Exercise 1: time series prediction

In this exercise we'll be looking at a [dataset of avocado prices](https://www.kaggle.com/neuromusic/avocado-prices). You should start by downloading the dataset. We will be forcasting average avocado price for the next week. 

```python
df = pd.read_csv("data/avocado.csv", parse_dates=["Date"], index_col=0)
df.head()
```

Output:
```
Date  AveragePrice  Total Volume     4046       4225    4770  \
0 2015-12-27          1.33      64236.62  1036.74   54454.85   48.16   
1 2015-12-20          1.35      54876.98   674.28   44638.81   58.33   
2 2015-12-13          0.93     118220.22   794.70  109149.67  130.50   
3 2015-12-06          1.08      78992.15  1132.00   71976.41   72.58   
4 2015-11-29          1.28      51039.60   941.48   43838.39   75.78   

   Total Bags  Small Bags  Large Bags  XLarge Bags          type  year  region  
0     8696.87     8603.62       93.25          0.0  conventional  2015  Albany  
1     9505.56     9408.07       97.49          0.0  conventional  2015  Albany  
2     8145.35     8042.21      103.14          0.0  conventional  2015  Albany  
3     5811.16     5677.40      133.76          0.0  conventional  2015  Albany  
4     6183.95     5986.26      197.69          0.0  conventional  2015  Albany
```

```python
df.shape
```

Output:
```
(18249, 13)
```

```python
df["Date"].min()
```

Output:
```
Timestamp('2015-01-04 00:00:00')
```

```python
df["Date"].max()
```

Output:
```
Timestamp('2018-03-25 00:00:00')
```

It looks like the data ranges from the start of 2015 to March 2018 (~2 years ago), for a total of 3.25 years or so. Let's split the data so that we have a 6 months of test data.

```python
split_date = '20170925'
df_train = df[df["Date"] <= split_date]
df_test  = df[df["Date"] >  split_date]
```

```python
assert len(df_train) + len(df_test) == len(df)
```

<br><br>

<!-- BEGIN QUESTION -->

### 1.1 How many time series? 
rubric={points:4}

In the [Rain in Australia](https://www.kaggle.com/datasets/jsphyg/weather-dataset-rattle-package) dataset from lecture, we had different measurements for each Location. How about this avocado sales dataset? For which categorical feature(s), if any, do we have separate measurements? Justify your answer by referencing the dataset.

<div class="alert alert-warning">

Solution_1.1
    
</div>

_Points:_ 4

In this dataset, the categorical feature "region" represents separate measurements. Each row in the dataset corresponds to a specific region where avocados are sold.  
If we analyze the dataset, we can observe that the "region" column contains different values, indicating various geographical locations such as Albany, Atlanta, Baltimore, and so on. Each region has its own set of avocado sales data, including average price, total volume, and other related features. Therefore, the "region" feature provides separate measurements for different locations or regions.

```python
categorical_features = ["type", "year", "region"]

for feature in categorical_features:
    unique_values = df[feature].unique()
    print(f"Feature: {feature}")
    print(f"Unique Values: {unique_values}")
    print(f"Count: {len(unique_values)}\n")

```

Output:
```
Feature: type
Unique Values: ['conventional' 'organic']
Count: 2

Feature: year
Unique Values: [2015 2016 2017 2018]
Count: 4

Feature: region
Unique Values: ['Albany' 'Atlanta' 'BaltimoreWashington' 'Boise' 'Boston'
 'BuffaloRochester' 'California' 'Charlotte' 'Chicago' 'CincinnatiDayton'
 'Columbus' 'DallasFtWorth' 'Denver' 'Detroit' 'GrandRapids' 'GreatLakes'
 'HarrisburgScranton' 'HartfordSpringfield' 'Houston' 'Indianapolis'
 'Jacksonville' 'LasVegas' 'LosAngeles' 'Louisville' 'MiamiFtLauderdale'
 'Midsouth' 'Nashville' 'NewOrleansMobile' 'NewYork' 'Northeast'
 'NorthernNewEngland' 'Orlando' 'Philadelphia' 'PhoenixTucson'
 'Pittsburgh' 'Plains' 'Portland' 'RaleighGreensboro' 'RichmondNorfolk'
 'Roanoke' 'Sacramento' 'SanDiego' 'SanFrancisco' 'Seattle'
 'SouthCarolina' 'SouthCentral' 'Southeast' 'Spokane' 'StLouis' 'Syracuse'
 'Tampa' 'TotalUS' 'West' 'WestTexNewMexico']
Count: 54
```

```python
...
```

Output:
```
Ellipsis
```

```python
...
```

Output:
```
Ellipsis
```

<!-- END QUESTION -->

<br><br>

<!-- BEGIN QUESTION -->

### 1.2 Equally spaced measurements? 
rubric={points:4}

In the Rain in Australia dataset, the measurements were generally equally spaced but with some exceptions. How about with this dataset? Justify your answer by referencing the dataset.

<div class="alert alert-warning">

Solution_1.2
    
</div>

_Points:_ 4

To determine if the measurements in the avocado sales dataset are equally spaced, we can calculate the time intervals between consecutive dates in the dataset. If the intervals are consistent, it suggests that the measurements are equally spaced. By calculating the differences between consecutive dates (date_diffs), we can check if all the differences are equal (is_equally_spaced). Since is_equally_spaced is False, it indicates that the measurements in the dataset are not equally spaced.

```python
date_diffs = df["Date"].diff()

is_equally_spaced = all(date_diffs == date_diffs.iloc[1])

print("Measurements in the avocado sales dataset are equally spaced:", is_equally_spaced)
```

Output:
```
Measurements in the avocado sales dataset are equally spaced: False
```

To check if the measurements in the avocado sales dataset are generally equally spaced but with some exceptions, we can calculate the time intervals between consecutive dates and analyze the distribution of these intervals. By calculating the differences between consecutive dates (date_diffs), we can determine the mean (mean_diff) and standard deviation (std_diff) of these differences. Additionally, we identify the maximum difference (max_diff) in the dataset. We then define a threshold for deviation from equal spacing, such as one standard deviation (threshold = mean_diff + std_diff). If any time interval exceeds the threshold, it indicates a deviation from equal spacing. Since has_exceptions is True, it means that there are exceptions to the general equal spacing of measurements in the dataset.

```python
date_diffs = df["Date"].diff()

mean_diff = date_diffs.mean()
std_diff = date_diffs.std()
max_diff = date_diffs.max()

threshold = mean_diff + std_diff

has_exceptions = any(date_diffs > threshold)

print("Measurements in the avocado sales dataset are generally equally spaced but with exceptions:", has_exceptions)
```

Output:
```
Measurements in the avocado sales dataset are generally equally spaced but with exceptions: True
```

```python
...
```

Output:
```
Ellipsis
```

```python
...
```

Output:
```
Ellipsis
```

```python
...
```

Output:
```
Ellipsis
```

```python
...
```

Output:
```
Ellipsis
```

<!-- END QUESTION -->

<br><br>

<!-- BEGIN QUESTION -->

### 1.3 Interpreting regions 
rubric={points:4}

In the Rain in Australia dataset, each location was a different place in Australia. For this dataset, look at the names of the regions. Do you think the regions are also all distinct, or are there overlapping regions? Justify your answer by referencing the data.

<div class="alert alert-warning">

Solution_1.3
    
</div>

_Points:_ 4

By looking at the names of the regions, they seem to all be states and cities in the USA. Since there are names of both cities and states, the regions are not distinct, and they do overlap as cities are inside a state. For example 'Los Angeles', 'San Francisco', 'Sacramento' and 'San Diego'  are all cities in the state 'California' so there is an overlap of regions there.

```python
regions = df['region'].unique()

print(regions)
```

Output:
```
['Albany' 'Atlanta' 'BaltimoreWashington' 'Boise' 'Boston'
 'BuffaloRochester' 'California' 'Charlotte' 'Chicago' 'CincinnatiDayton'
 'Columbus' 'DallasFtWorth' 'Denver' 'Detroit' 'GrandRapids' 'GreatLakes'
 'HarrisburgScranton' 'HartfordSpringfield' 'Houston' 'Indianapolis'
 'Jacksonville' 'LasVegas' 'LosAngeles' 'Louisville' 'MiamiFtLauderdale'
 'Midsouth' 'Nashville' 'NewOrleansMobile' 'NewYork' 'Northeast'
 'NorthernNewEngland' 'Orlando' 'Philadelphia' 'PhoenixTucson'
 'Pittsburgh' 'Plains' 'Portland' 'RaleighGreensboro' 'RichmondNorfolk'
 'Roanoke' 'Sacramento' 'SanDiego' 'SanFrancisco' 'Seattle'
 'SouthCarolina' 'SouthCentral' 'Southeast' 'Spokane' 'StLouis' 'Syracuse'
 'Tampa' 'TotalUS' 'West' 'WestTexNewMexico']
```

```python
...
```

Output:
```
Ellipsis
```

```python
...
```

Output:
```
Ellipsis
```

<!-- END QUESTION -->

<br><br>

We will use the entire dataset despite any location-based weirdness uncovered in the previous part.

We will be trying to forecast the avocado price. The function below is adapted from [Lecture 19](https://github.com/UBC-CS/cpsc330-2023s/blob/main/lectures/19_time-series.ipynb), with some improvements.

```python
def create_lag_feature(df, orig_feature, lag, groupby, new_feature_name=None, clip=False):
    """
    Creates a new feature that's a lagged version of an existing one.
    
    NOTE: assumes df is already sorted by the time columns and has unique indices.
    
    Parameters
    ----------
    df : pandas.core.frame.DataFrame
        The dataset.
    orig_feature : str
        The column name of the feature we're copying
    lag : int
        The lag; negative lag means values from the past, positive lag means values from the future
    groupby : list
        Column(s) to group by in case df contains multiple time series
    new_feature_name : str
        Override the default name of the newly created column
    clip : bool
        If True, remove rows with a NaN values for the new feature
    
    Returns
    -------
    pandas.core.frame.DataFrame
        A new dataframe with the additional column added.
        
    """
        
    if new_feature_name is None:
        if lag < 0:
            new_feature_name = "%s_lag%d" % (orig_feature, -lag)
        else:
            new_feature_name = "%s_ahead%d" % (orig_feature, lag)
    
    new_df = df.assign(**{new_feature_name : np.nan})
    for name, group in new_df.groupby(groupby):        
        if lag < 0: # take values from the past
            new_df.loc[group.index[-lag:],new_feature_name] = group.iloc[:lag][orig_feature].values
        else:       # take values from the future
            new_df.loc[group.index[:-lag], new_feature_name] = group.iloc[lag:][orig_feature].values
            
    if clip:
        new_df = new_df.dropna(subset=[new_feature_name])
        
    return new_df
```

We first sort our dataframe properly:

```python
df_sort = df.sort_values(by=["region", "type", "Date"]).reset_index(drop=True)
df_sort
```

Output:
```
Date  AveragePrice  Total Volume     4046      4225    4770  \
0     2015-01-04          1.22      40873.28  2819.50  28287.42   49.90   
1     2015-01-11          1.24      41195.08  1002.85  31640.34  127.12   
2     2015-01-18          1.17      44511.28   914.14  31540.32  135.77   
3     2015-01-25          1.06      45147.50   941.38  33196.16  164.14   
4     2015-02-01          0.99      70873.60  1353.90  60017.20  179.32   
...          ...           ...           ...      ...       ...     ...   
18244 2018-02-25          1.57      18421.24  1974.26   2482.65    0.00   
18245 2018-03-04          1.54      17393.30  1832.24   1905.57    0.00   
18246 2018-03-11          1.56      22128.42  2162.67   3194.25    8.93   
18247 2018-03-18          1.56      15896.38  2055.35   1499.55    0.00   
18248 2018-03-25          1.62      15303.40  2325.30   2171.66    0.00   

       Total Bags  Small Bags  Large Bags  XLarge Bags          type  year  \
0         9716.46     9186.93    
```

We then call `create_lag_feature`. This creates a new column in the dataset `AveragePriceNextWeek`, which is the following week's `AveragePrice`. We have set `clip=True` which means it will remove rows where the target would be missing.

```python
df_hastarget = create_lag_feature(df_sort, "AveragePrice", +1, ["region", "type"], "AveragePriceNextWeek", clip=True)
df_hastarget
```

Output:
```
Date  AveragePrice  Total Volume     4046      4225    4770  \
0     2015-01-04          1.22      40873.28  2819.50  28287.42   49.90   
1     2015-01-11          1.24      41195.08  1002.85  31640.34  127.12   
2     2015-01-18          1.17      44511.28   914.14  31540.32  135.77   
3     2015-01-25          1.06      45147.50   941.38  33196.16  164.14   
4     2015-02-01          0.99      70873.60  1353.90  60017.20  179.32   
...          ...           ...           ...      ...       ...     ...   
18243 2018-02-18          1.56      17597.12  1892.05   1928.36    0.00   
18244 2018-02-25          1.57      18421.24  1974.26   2482.65    0.00   
18245 2018-03-04          1.54      17393.30  1832.24   1905.57    0.00   
18246 2018-03-11          1.56      22128.42  2162.67   3194.25    8.93   
18247 2018-03-18          1.56      15896.38  2055.35   1499.55    0.00   

       Total Bags  Small Bags  Large Bags  XLarge Bags          type  year  \
0         9716.46     9186.93    
```

Our goal is to predict `AveragePriceNextWeek`. 

Let's split the data:

```python
df_train = df_hastarget[df_hastarget["Date"] <= split_date]
df_test  = df_hastarget[df_hastarget["Date"] >  split_date]
```

<br><br>

### 1.4 `AveragePrice` baseline 
rubric={autograde:4}

Soon we will want to build some models to forecast the average avocado price a week in advance. Before we start with any ML though, let's try a baseline. Previously we used `DummyClassifier` or `DummyRegressor` as a baseline. This time, we'll do something else as a baseline: we'll assume the price stays the same from this week to next week. So, we'll set our prediction of "AveragePriceNextWeek" exactly equal to "AveragePrice", assuming no change. That is kind of like saying, "If it's raining today then I'm guessing it will be raining tomorrow". This simplistic approach will not get a great score but it's a good starting point for reference. If our model does worse that this, it must not be very good. 

Using this baseline approach, what $R^2$ do you get on the train and test data?

<div class="alert alert-warning">

Solution_1.4
    
</div>

_Points:_ 4

_Type your answer here, replacing this text._

```python
train_r2 = None

train_r2 = r2_score(df_train['AveragePriceNextWeek'], df_train['AveragePrice'])
```

```python
test_r2 = None

test_r2 = r2_score(df_test['AveragePriceNextWeek'], df_test['AveragePrice'])
```

```python
...
```

Output:
```
Ellipsis
```

```python
...
```

Output:
```
Ellipsis
```

```python
grader.check("q1.4")
```

Output:
```
q1.4 results: All test cases passed!
```

<br><br>

<!-- BEGIN QUESTION -->

### 1.5 Forecasting average avocado price
rubric={points:10}

Now that the baseline is done, let's build some models to forecast the average avocado price a week later. Experiment with a few approachs for encoding the date. Justify the decisions you make. Which approach worked best? Report your test score and briefly discuss your results.

Benchmark: you should be able to achieve $R^2$ of at least 0.79 on the test set. I got to 0.80, but not beyond that. Let me know if you do better!

Note: because we only have 2 splits here, we need to be a bit wary of overfitting on the test set. Try not to test on it a ridiculous number of times. If you are interested in some proper ways of dealing with this, see for example sklearn's [TimeSeriesSplit](https://scikit-learn.org/stable/modules/generated/sklearn.model_selection.TimeSeriesSplit.html), which is like cross-validation for time series data.

<div class="alert alert-warning">

Solution_1.3
    
</div>

_Points:_ 10

_Type your answer here, replacing this text._

```python

# Encode ordinal features: day of the week and month
df_sort['Date'] = pd.to_datetime(df_sort['Date'])
df_sort['DayOfWeek'] = df_sort['Date'].dt.dayofweek
df_sort['Month'] = df_sort['Date'].dt.month

# Perform ordinal encoding for the 'region' column
ordinal_encoder = OrdinalEncoder()
df_sort['region_code'] = ordinal_encoder.fit_transform(df_sort[['region']])

# Prepare feature matrix and target variable
features = ['DayOfWeek', 'Month', 'Total Volume', '4046', '4225', '4770', 'Total Bags', 'Small Bags', 'Large Bags', 'XLarge Bags', 'year', 'region_code']
X_train = df_sort.loc[df_sort['Date'] <= split_date, features]
X_test = df_sort.loc[df_sort['Date'] > split_date, features]
y_train = df_sort.loc[df_sort['Date'] <= split_date, 'AveragePrice']
y_test = df_sort.loc[df_sort['Date'] > split_date, 'AveragePrice']


rf_model = RandomForestRegressor(n_estimators=100, random_state=42)
rf_model.fit(X_train, y_train)


y_pred = rf_model.predict(X_test)


test_r2 = r2_score(y_test, y_pred)
print(test_r2)

```

Output:
```
0.2886923437649255
```

```python

```

```python
...
```

Output:
```
Ellipsis
```

```python
...
```

Output:
```
Ellipsis
```

```python
...
```

Output:
```
Ellipsis
```

```python
...
```

Output:
```
Ellipsis
```

```python
...
```

Output:
```
Ellipsis
```

```python
...
```

Output:
```
Ellipsis
```

```python
...
```

Output:
```
Ellipsis
```

```python
...
```

Output:
```
Ellipsis
```

```python
...
```

Output:
```
Ellipsis
```

```python
...
```

Output:
```
Ellipsis
```

```python
...
```

Output:
```
Ellipsis
```

```python
...
```

Output:
```
Ellipsis
```

```python
...
```

Output:
```
Ellipsis
```

```python
...
```

Output:
```
Ellipsis
```

```python
...
```

Output:
```
Ellipsis
```

```python
...
```

Output:
```
Ellipsis
```

```python
...
```

Output:
```
Ellipsis
```

```python
...
```

Output:
```
Ellipsis
```

```python
...
```

Output:
```
Ellipsis
```

```python
...
```

Output:
```
Ellipsis
```

```python
...
```

Output:
```
Ellipsis
```

```python
...
```

Output:
```
Ellipsis
```

<!-- END QUESTION -->

<br><br><br><br>

## Exercise 2: very short answer questions

Each question is worth 2 points.

<!-- BEGIN QUESTION -->

### 2.1 Time series

rubric={points:4}

The following questions pertain to Lecture 19 on time series data:

1. Sometimes a time series has missing time points or, worse, time points that are unequally spaced in general. Give an example of a real world situation where the time series data would have unequally spaced time points.
2. In class we discussed two approaches to using temporal information: encoding the date as one or more features, and creating lagged versions of features. Which of these (one/other/both/neither) two approaches would struggle with unequally spaced time points? Briefly justify your answer.

<div class="alert alert-warning">

Solution_2.1
    
</div>

_Points:_ 4

1. An example of a real-world situation where time series data would have unequally spaced time points is weather data. Weather observations are typically recorded at specific intervals, such as every hour or every few hours. However, due to various factors like technical issues, maintenance, or extreme weather conditions, some observations may be missed or recorded at irregular intervals. This can result in unequally spaced time points in the weather data.

2. When dealing with unequally spaced time points, encoding the date as one or more features would struggle. Encoding the date as a feature assumes a regular and consistent time interval between data points. However, with unequally spaced time points, this assumption is violated, and encoding the date would not capture the irregularity in the time series.

<!-- END QUESTION -->

<br><br>

<!-- BEGIN QUESTION -->

### 2.2 Survival analysis
rubric={points:6}

The following questions pertain to [Lecture 20](https://github.com/UBC-CS/cpsc330-2023s/blob/main/lectures/20_survival-analysis.ipynb) on survival analysis. We'll consider the use case of customer churn analysis.

1. What is the problem with simply labeling customers are "churned" or "not churned" and using standard supervised learning techniques?
2. Consider customer A who just joined last week vs. customer B who has been with the service for a year. Who do you expect will leave the service first: probably customer A, probably customer B, or we don't have enough information to answer?
3. If a customer's survival function is almost flat during a certain period, how do we interpret that?

<div class="alert alert-warning">

Solution_2.2
    
</div>

_Points:_ 6

1. The problem is that it ignores the temporal aspect of churn. In customer churn analysis, the timing or duration of customer survival is crucial information. Standard supervised learning techniques treat each observation independently and assume that the timing of events is irrelevant. However, in churn analysis, the duration of a customer's relationship with a service can significantly impact the likelihood of churn. Ignoring the time dimension can lead to biased and inaccurate predictions.

2. Without further information, it is difficult to determine whether customer A or customer B is more likely to leave the service first. The duration of a customer's relationship with the service is one factor to consider, but other individual characteristics, behaviors, or external factors may also influence churn. Customer A, who just joined last week, may have specific reasons or circumstances that make them more likely to churn early. On the other hand, customer B, who has been with the service for a year, may have developed loyalty or engagement that reduces their likelihood of churn. Additional information about the customers' behaviors, preferences, or interactions with the service would be necessary to make a more accurate prediction.

3. If a customer's survival function is almost flat during a certain period, it indicates that the probability of the customer churning is relatively constant or stagnant over that time period. In other words, the customer's likelihood of churning remains relatively unchanged. This could suggest that the customer is exhibiting a stable behavior or that there are external factors influencing their churn probability that are not captured in the analysis. It is important to further investigate and consider additional factors that may be affecting the customer's churn probability to gain a deeper understanding of their behavior and potential risks of churn.

<!-- END QUESTION -->

<br><br>

**PLEASE READ BEFORE YOU SUBMIT:** 

When you are ready to submit your assignment do the following:

1. Run all cells in your notebook to make sure there are no errors by doing `Kernel -> Restart Kernel and Clear All Outputs` and then `Run -> Run All Cells`. 
2. Notebooks with cell execution numbers out of order or not starting from "1" will have marks deducted. Notebooks without the output displayed may not be graded at all (because we need to see the output in order to grade your work).
3. Upload the assignment using Gradescope's drag and drop tool. Check out this [Gradescope Student Guide](https://lthub.ubc.ca/guides/gradescope-student-guide/) if you need help with Gradescope submission. 
4. Make sure that the plots and output are rendered properly in your submitted file. If the .ipynb file is too big and doesn't render on Gradescope, also upload a pdf or html in addition to the .ipynb so that the TAs can view your submission on Gradescope. 

![](img/eva-well-done.png)

Tests for q1.4 (points: 4)
```
>>> assert not train_r2 is None, "Are you using the correct variable name?"
>>> assert not test_r2 is None, "Are you using the correct variable name?"
>>> assert sha1(str(round(train_r2, 3)).encode('utf8')).hexdigest() == 'b1136fe2a8918904393ab6f40bfb3f38eac5fc39', "Your training score is not correct. Are you using the right features?"
>>> assert sha1(str(round(test_r2, 3)).encode('utf8')).hexdigest() == 'cc24d9a9b567b491a56b42f7adc582f2eefa5907', "Your test score is not correct. Are you using the right features?"
```
